
<br>
<br>
<div>
<img  style="float: left; padding-right: 100px; width: 350px" src="./logo.png">
    </div>
    <br>
    <h3 align="center">AI4D LAB TRAINING</h3>
<hr>
<br>
<h4 align="center"><a href="https://nsoma.me">Zephania Reuben</a></h4>
<br>
<h4 align="center">July 18, 2023</h4>
<br>
<hr>
<h3 align="center">DATA SCIENCE | TASK</h3>
<hr>
<br>

<h3 align="center"> FINANCIAL INCLUSSION IN AFRICA - ZINDI COMPETITION </h3>

### Understand The Problem Statement

Financial Inclusion remains one of the main obstacles to economic and human development in Africa. For example, across Kenya, Rwanda, Tanzania, and Uganda only 9.1 million adults (or 13.9% of the adult population) have access to or use a commercial bank account.

Traditionally, access to bank accounts has been regarded as an indicator of financial inclusion. Despite the proliferation of mobile money in Africa, and the growth of innovative fintech solutions, banks still play a pivotal role in facilitating access to financial services. Access to bank accounts enable households to save and facilitate payments while also helping businesses build up their credit-worthiness and improve their access to other finance services. Therefore, access to bank accounts is an essential contributor to long-term economic growth.

The objective of this Dataset is to create a machine learning model to predict which individuals are most likely to have or use a bank account. The models and solutions developed can provide an indication of the state of financial inclusion in Kenya, Rwanda, Tanzania and Uganda, while providing insights into some of the key demographic factors that might drive individuals’ financial outcomes.

Data source available in the the zindi platform, [Zindi Africa](https://zindi.africa/competitions/financial-inclusion-in-africa)

### Hypothesis Generation
This is a very important stage in any data science/machine learning pipeline. It involves understanding the problem in detail by brainstorming as many factors as possible which can impact the outcome. It is done by understanding the problem statement thoroughly and before looking at the data.

Below are some of the factors which I think can affect the chance for a person to have a bank account

- People who have mobile phone have lower chance to use bank account because of mobile money services.
- People who are employed have a higher chance of having a bank account than People who are unemployed.
- People with low education have low chance to have bank account
- People in rural areas have low chance to have bank account 
- People who have age below 18 have low chance to have bank account
- Female have less chance to have bank account

```python
#Import libraries
import pandas as pd
import numpy as np
import seaborn as sns
import matplotlib.pyplot as plt
from sklearn.model_selection import train_test_split
from sklearn.linear_model import LogisticRegression
from sklearn.neighbors import KNeighborsClassifier
from sklearn.ensemble import RandomForestClassifier
from sklearn.tree import DecisionTreeClassifier
from sklearn.metrics import mean_squared_error
from sklearn.metrics import accuracy_score
```

```python
# Load files into a pandas dataframe
train = pd.read_csv("/home/jok/4Legends/Train.csv")
```

Credit: [Davis David](https://twitter.com/Davis_McDavid)

```python
train.head()
```

Output:
```
country  year    uniqueid bank_account location_type cellphone_access  \
0   Kenya  2018  uniqueid_1          Yes         Rural              Yes   
1   Kenya  2018  uniqueid_2           No         Rural               No   
2   Kenya  2018  uniqueid_3          Yes         Urban              Yes   
3   Kenya  2018  uniqueid_4           No         Rural              Yes   
4   Kenya  2018  uniqueid_5           No         Urban               No   

   household_size  age_of_respondent gender_of_respondent  \
0               3                 24               Female   
1               5                 70               Female   
2               5                 26                 Male   
3               5                 34               Female   
4               8                 26                 Male   

  relationship_with_head           marital_status  \
0                 Spouse  Married/Living together   
1      Head of Household                  Widowed   
2         Other relative 
```

```python
# shows shapes of the datasets
print(train.shape)
```

Output:
```
(23524, 13)
```

```python
train.head(6)
```

Output:
```
country  year    uniqueid bank_account location_type cellphone_access  \
0   Kenya  2018  uniqueid_1          Yes         Rural              Yes   
1   Kenya  2018  uniqueid_2           No         Rural               No   
2   Kenya  2018  uniqueid_3          Yes         Urban              Yes   
3   Kenya  2018  uniqueid_4           No         Rural              Yes   
4   Kenya  2018  uniqueid_5           No         Urban               No   
5   Kenya  2018  uniqueid_6           No         Rural               No   

   household_size  age_of_respondent gender_of_respondent  \
0               3                 24               Female   
1               5                 70               Female   
2               5                 26                 Male   
3               5                 34               Female   
4               8                 26                 Male   
5               7                 26               Female   

  relationship_with_head           marital_statu
```

```python
#Information about the columns in the data
train.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 23524 entries, 0 to 23523
Data columns (total 13 columns):
 #   Column                  Non-Null Count  Dtype 
---  ------                  --------------  ----- 
 0   country                 23524 non-null  object
 1   year                    23524 non-null  int64 
 2   uniqueid                23524 non-null  object
 3   bank_account            23524 non-null  object
 4   location_type           23524 non-null  object
 5   cellphone_access        23524 non-null  object
 6   household_size          23524 non-null  int64 
 7   age_of_respondent       23524 non-null  int64 
 8   gender_of_respondent    23524 non-null  object
 9   relationship_with_head  23524 non-null  object
 10  marital_status          23524 non-null  object
 11  education_level         23524 non-null  object
 12  job_type                23524 non-null  object
dtypes: int64(3), object(10)
memory usage: 2.3+ MB
```

```python
#Changing categorical data into numerical data
train['country'].replace(['Kenya','Rwanda','Tanzania','Uganda'],[0,1,2,3],inplace=True)
train['bank_account'].replace(['No','Yes'],[0,1],inplace=True)
train['location_type'].replace(['Rural','Urban'],[0,1],inplace=True)
train['cellphone_access'].replace(['No','Yes'],[0,1],inplace=True)
train['gender_of_respondent'].replace(['Female','Male'],[0,1],inplace=True)
train['relationship_with_head'].replace(['Spouse','Head of Household','Other relative','Child','Parent','Other non-relatives'],[0,1,2,3,4,5],inplace=True)
train['marital_status'].replace(['Married/Living together','Widowed','Single/Never Married','Divorced/Seperated','Dont know'],[0,1,2,3,4],inplace=True)
train['education_level'].replace(['Secondary education','No formal education','Vocational/Specialised training','Primary education','Tertiary education','Other/Dont know/RTA'],[0,1,2,3,4,5],inplace=True)
train['job_type'].replace(['Self employed','Government Dependent','Formally employed Private','Informally employed','Formally employed Government','Farming and Fishing','Remittance Dependent','Other Income','Dont Know/Refuse to answer','No Income'],[0,1,2,3,4,5,6,7,8,9],inplace=True)
```

```python
#Dropping uniqueid columns
trainee = train.drop(['uniqueid'],axis = 1)
```

```python
trainee.head(10)
```

Output:
```
country  year  bank_account  location_type  cellphone_access  \
0        0  2018             1              0                 1   
1        0  2018             0              0                 0   
2        0  2018             1              1                 1   
3        0  2018             0              0                 1   
4        0  2018             0              1                 0   
5        0  2018             0              0                 0   
6        0  2018             0              0                 1   
7        0  2018             0              0                 1   
8        0  2018             1              0                 1   
9        0  2018             0              1                 1   

   household_size  age_of_respondent  gender_of_respondent  \
0               3                 24                     0   
1               5                 70                     0   
2               5                 26                     1   
3               5
```

```python
X = trainee.drop('bank_account', axis=1)
y = trainee['bank_account']

# Split the data into training and testing sets
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=42)
```

```python
#Training the models using train data set
#logisticRegression
model=LogisticRegression()
model.fit(X_train, y_train)

# K-Nearest Neighbors (KNN)
knn_model = KNeighborsClassifier(n_neighbors=5)
knn_model.fit(X_train, y_train)

# Random Forest
rf_model = RandomForestClassifier(n_estimators=100, random_state=42)
rf_model.fit(X_train, y_train)

# Decision Tree
dt_model = DecisionTreeClassifier(random_state=42)
dt_model.fit(X_train, y_train)
```

Output:
```
/home/jok/.local/lib/python3.10/site-packages/sklearn/linear_model/_logistic.py:460: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-learn.org/stable/modules/linear_model.html#logistic-regression
  n_iter_i = _check_optimize_result(
DecisionTreeClassifier(random_state=42)
```

```python
# Predict using each model
log_pred=model.predict(X_test)
knn_pred = knn_model.predict(X_test)
rf_pred = rf_model.predict(X_test)
dt_pred = dt_model.predict(X_test)
```

```python
#Showing the result of predictions
print('Logistic Prediction',log_pred)
print('KNN Prediction',knn_pred)
print('Random Forest Prediction',rf_pred)
print('Decision Tree Prediction',dt_pred)
```

Output:
```
Logistic Prediction [0 0 0 ... 0 0 0]
KNN Prediction [0 0 0 ... 0 0 0]
Random Forest Prediction [0 0 1 ... 0 0 0]
Decision Tree Prediction [0 0 1 ... 0 0 0]
```

```python
#Evaluating the performance and outputing their performance
meansquared_log=mean_squared_error(y_test,log_pred)
print('Performance of Logistic Regresssion',meansquared_log)
meansquared_knn=mean_squared_error(y_test,knn_pred)
print('Performance of KNN',meansquared_knn)
meansquared_rf=mean_squared_error(y_test,rf_pred)
print('Performance of Random Forest',meansquared_rf)
meansquared_dt=mean_squared_error(y_test,dt_pred)
print('Performance of Decision Tree',meansquared_dt)
```

Output:
```
Performance of Logistic Regresssion 0.13262486716259297
Performance of KNN 0.14027630180658873
Performance of Random Forest 0.13751328374070138
Performance of Decision Tree 0.15855472901168968
```

```python
#Importing test data set
test = pd.read_csv("/home/jok/4Legends/Test.csv")
```

```python
test.head()
```

Output:
```
country  year       uniqueid location_type cellphone_access  household_size  \
0   Kenya  2018  uniqueid_6056         Urban              Yes               3   
1   Kenya  2018  uniqueid_6060         Urban              Yes               7   
2   Kenya  2018  uniqueid_6065         Rural               No               3   
3   Kenya  2018  uniqueid_6072         Rural               No               6   
4   Kenya  2018  uniqueid_6073         Urban               No               3   

   age_of_respondent gender_of_respondent relationship_with_head  \
0                 30                 Male      Head of Household   
1                 51                 Male      Head of Household   
2                 77               Female                 Parent   
3                 39               Female      Head of Household   
4                 16                 Male                  Child   

            marital_status                  education_level  \
0  Married/Living together              Sec
```

```python
#Changing categorical data into numerical data
test['country'].replace(['Kenya','Rwanda','Tanzania','Uganda'],[0,1,2,3],inplace=True)
test['location_type'].replace(['Rural','Urban'],[0,1],inplace=True)
test['cellphone_access'].replace(['No','Yes'],[0,1],inplace=True)
test['gender_of_respondent'].replace(['Female','Male'],[0,1],inplace=True)
test['relationship_with_head'].replace(['Spouse','Head of Household','Other relative','Child','Parent','Other non-relatives'],[0,1,2,3,4,5],inplace=True)
test['marital_status'].replace(['Married/Living together','Widowed','Single/Never Married','Divorced/Seperated','Dont know'],[0,1,2,3,4],inplace=True)
test['education_level'].replace(['Secondary education','No formal education','Vocational/Specialised training','Primary education','Tertiary education','Other/Dont know/RTA'],[0,1,2,3,4,5],inplace=True)
test['job_type'].replace(['Self employed','Government Dependent','Formally employed Private','Informally employed','Formally employed Government','Farming and Fishing','Remittance Dependent','Other Income','Dont Know/Refuse to answer','No Income'],[0,1,2,3,4,5,6,7,8,9],inplace=True)
```

```python
#Dropping uniqueid columns
tester = test.drop(['uniqueid'],axis = 1)
```

```python
tester
```

Output:
```
country  year  location_type  cellphone_access  household_size  \
0            0  2018              1                 1               3   
1            0  2018              1                 1               7   
2            0  2018              0                 0               3   
3            0  2018              0                 0               6   
4            0  2018              1                 0               3   
...        ...   ...            ...               ...             ...   
10081        3  2018              0                 0               2   
10082        3  2018              1                 1               8   
10083        3  2018              1                 1               1   
10084        3  2018              0                 1               6   
10085        3  2018              1                 1               7   

       age_of_respondent  gender_of_respondent  relationship_with_head  \
0                     30                     1          
```

```python
#Predicting the valuse using test data set
log_pred=model.predict(tester)
knn_pred = knn_model.predict(tester)
rf_pred = rf_model.predict(tester)
dt_pred = dt_model.predict(tester)

```

```python
#Showing the result of predictions
print('Logistic Prediction',log_pred)
print('KNN Prediction',knn_pred)
print('Random Forest Prediction',rf_pred)
print('Decision Tree Prediction',dt_pred)
```

Output:
```
Logistic Prediction [0 0 0 ... 0 0 0]
KNN Prediction [1 1 0 ... 0 0 0]
Random Forest Prediction [1 1 0 ... 0 0 0]
Decision Tree Prediction [1 1 0 ... 0 0 0]
```

```python
#Assigning prediction to variable mypredict
mypredict = dt_pred
```

```python
test["bank_account"]= mypredict
```

```python
#Changing the numerical data into categorical data 
test['country'].replace([0,1,2,3],['Kenya','Rwanda','Tanzania','Uganda'],inplace=True)
```

```python
test.head()
```

Output:
```
country  year       uniqueid  location_type  cellphone_access  \
0   Kenya  2018  uniqueid_6056              1                 1   
1   Kenya  2018  uniqueid_6060              1                 1   
2   Kenya  2018  uniqueid_6065              0                 0   
3   Kenya  2018  uniqueid_6072              0                 0   
4   Kenya  2018  uniqueid_6073              1                 0   

   household_size  age_of_respondent  gender_of_respondent  \
0               3                 30                     1   
1               7                 51                     1   
2               3                 77                     0   
3               6                 39                     0   
4               3                 16                     1   

   relationship_with_head  marital_status  education_level  job_type  \
0                       1               0                0         4   
1                       1               0                2         2   
2         
```

```python
#Joining data for zindi
ready_for_zindi = test[['uniqueid','country','bank_account']]
```

```python
#Displaying the data after joining columns needed by zindi
ready_for_zindi
```

Output:
```
uniqueid country  bank_account
0      uniqueid_6056   Kenya             1
1      uniqueid_6060   Kenya             1
2      uniqueid_6065   Kenya             0
3      uniqueid_6072   Kenya             0
4      uniqueid_6073   Kenya             0
...              ...     ...           ...
10081  uniqueid_2998  Uganda             0
10082  uniqueid_2999  Uganda             0
10083  uniqueid_3000  Uganda             0
10084  uniqueid_3001  Uganda             0
10085  uniqueid_3002  Uganda             0

[10086 rows x 3 columns]
```

```python
#Exporting data to csv
ready_for_zindi.to_csv("/home/jok/4Legends/zindii.csv",sep=',',index=False,encoding = 'utf-8')
```

Tests for q1 (points: 6)
```
>>> ypoly(3,3) == 139
True
```
```
>>> ypoly(-2,3) == 0
True
```


Tests for q2 (points: 5)
```
>>> assert_equal(perfectly_divisible(10,2), [0,2,4,6,8,10])
```


Tests for q3 (points: 4)
```
>>> assert_equal(flatten_lists([[2,13,44], [6,7]]), [2, 6, 7, 13, 44])
```


Tests for q4 (points: 5)
```
>>> assert_equal(longest_word("Hello there people"), 'people')
```


Tests for q5 (points: 5)
```
>>> assert_equal(dictionarize([1,2,2,3,4,5]), {1: 1, 2: 2, 3: 1, 4: 1, 5: 1})
```


Tests for q6 (points: 5)
```
>>> assert_equal(greatest_common_divisor(10, 15), 5)
```


Tests for q7 (points: 10)
```
>>> assert_equal(get_nearest_farthest((3, 8), [(9, 3), (8,5), (7,6)]), ((7, 6), (9, 3)))
```
